Noms et prénoms du binome :

<!-- grade_id: cell-7b535f8957fa16ad, points: 1, grade: true, solution: true -->


## Polynômes du second degré

Nous allons travailler dans cette activité sur l'utilisation des algorithmes et de la programmation pour résoudre des équations du second degré.

Dans ce document, les cellules de code sont à exécuter soit à l'aide du bouton Exécuter ci-dessus, soit à l'aide des combinaisons de touches Ctrl + Entrée. Les cellules de code sont librement modifiable, il faut ensuite les ré-exécuter.

Les `# A REMPLACER` sont à remplacer par les instructions Python adéquates.

### 1) Exercice
Pour commencer, nous allons compléter le programme Python suivant qui demande les coefficients d'un polynôme du second degré $ax^2+bx+c$ et qui annonce le nombre de racines dupolynôme et les affiche s'il y en a.

```python
from math import sqrt

a = float(input('Saisir la valeur du paramètre a :'))
b = float(input('Saisir la valeur du paramètre b :'))
c = float(input('Saisir la valeur du paramètre c :'))

delta = # A REMPLACER

if delta == 0:
    r1 = # A REMPLACER
    print('...')
elif delta > 0 :
    r1 = # A REMPLACER
    r2 = # A REMPLACER
    print('...')
else :
    print('...')
```

1.1) Compléter ce programme afin qu'il fonctionne

_Note:_ l'instruction `print` peut prendre plusieurs arguments, par exemple  `print('afficher', a, 'et', b) `

```python
print('afficher', 5, 'et', 2)
```

1.2) Sur la première ligne, à quoi sert l'instruction `input` ? Et l'instruction `float` ?

<!-- grade_id: cell-bb42641fd934eb83, points: 2, grade: true, solution: true -->


1.3) Sur la 7e ligne démarrant par `if`, pourquoi y a-t-il deux signes égal (`==`) et non pas un seul ?

<!-- grade_id: cell-80fe44fc28ed7bf4, points: 1, grade: true, solution: true -->


### 2) Création d'une fonction renvoyant le nombre de solutions d'un polynome
On va maintenant modifier le programme afin qu'un fonction se charge du calcul du nombre de solutions.

2.1) On va donc définir une fonction `nb_sol_polynome` qui prendra en paramètre les arguments `a`, `b` et `c`. On utilisera l'instruction `return` pour renvoyer la valeur de retour.

```python
def nb_sol_polynome(a, b, c):
    assert(a != 0)
    
    delta = # A REMPLACER
    
    if ## A REMPLACER :
        return # A REMPLACER
    # A REMPLACER
```

Ensuite on pourra appeler cette fonction pour obtenir le nombre de solution.

```python
nb_sol_polynome(2, 1, 5)
```

```python
nb_sol_polynome(0, 1, 5)
```

2.2) A quoi sert l'instruction `assert(a != 0)` au début de la fonction ?

<!-- grade_id: cell-50852c519fd5a416, points: 1, grade: true, solution: true -->


On peut aussi faire vérifier automatiquement à Python que la fonction parait fonctionner comme on le souhaite en testant quelques exemples. 

**Python va générer une erreur uniquement si une des fonction n'a pas renvoyé le résultat attendu.** Dans ce cas, corriger votre fonction afin que le résultat renvoyé soit correct.

<!-- grade_id: cell-d55cde5b4b5bef2d, points: 6, grade: true, solution: false -->
```python
assert nb_sol_polynome(2, -1, 5) == 0
assert nb_sol_polynome(-2, 5, 1) == 2
assert nb_sol_polynome(1, 0, 0) == 1
### BEGIN HIDDEN TESTS
assert nb_sol_polynome(1, 1, -1) == 2
assert nb_sol_polynome(1, -3, 4) == 0
assert nb_sol_polynome(1, -2, 1) == 1
### END HIDDEN TESTS
```

### 3) Création d'une fonction qui renvoie les solutions approchées d'un polynome

3.1) On cherche maintenant à créer une nouvelle fonction `solutions_polynome` qui prend également 3 arguments en paramètre et qui renvoie les solutions sous forme de tableau, et rangées par ordre croissant.

_Note :_ Un tableau de valeur est défini avec des crochets et en utilisant la virgule pour séparer les éléments. Par exemple, `[-1.5, 5]`, `[0.33333333]` ou `[]`

```python
def solutions_polynome(a, b, c):
    # A REMPLACER

```

```python
solutions_polynome(-2, 5, 1)
```

<!-- grade_id: cell-0c814118388abe68, points: 6, grade: true, solution: false -->
```python
assert solutions_polynome(-2, 5, 1) == [-0.18614066163450715, 2.686140661634507]
assert solutions_polynome(-2, -5, 1) == [-2.686140661634507, 0.18614066163450715]
assert solutions_polynome(1, 0, 0) == [0.0]
### BEGIN HIDDEN TESTS
assert solutions_polynome(1, 1, -1) == [-1.618033988749895, 0.6180339887498949]
assert solutions_polynome(1, -3, 4) == []
assert solutions_polynome(1, -2, 1) == [1.0]
### END HIDDEN TESTS
```